---
title: "Tutorial"
output: learnr::tutorial
runtime: shiny_prerendered
description: "Introduce interactive tutorials"
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
tutorial_options(exercise.checker = "gradethis::grade_learnr()")
knitr::opts_chunk$set(echo = FALSE)
```


# Overview  
  
This document is aimed to introduce the possibilities of an interactive tutorial using `rmarkdown` and the `learnr` package. We can embed code chunks from both r and python, making this format extremely versatile.   
  
  
## Media options  
  
Let's take a look at the media options we can use within an interactive tutorials.   
  
**Videos**  
  
We can easily add a video to our file by embedding them in our rmarkdown file:  
![](https://www.youtube.com/watch?v=YddwkMJG1Jo)  
  
This will facilitate many of our tutorials as we can just link to a video on youtube or vimeo for a short introduction to the specific topic.  
  
  
**Images**  
  
The same holds for images and such. We can link to images hosted on our server or from the web:  
  
```{r image, out.width = '50%', out.height= '30%', fig.align="center"}
knitr::include_graphics("https://i.redd.it/6jla414oko0z.jpg")
```  
  
  
**Links**  
  
We can, of cause, link to all sources that are available on the web. For example, have you heared about this weird guy doing [stats with extinct critters](https://gregor-mathes.netlify.app/)?  
  
## Code chunks  
  
The greatest advantage of an interactive tutorial compared to a static html is that one can create exercises with direct evaluation. This is extremely customizable and powerfull (*e.g. tutorials store the various pieces of work done within them such as exercise input and output, answers to questions, etc., and automatically restore it when users revisit them. Locally run tutorials store user work in the local filesystem and tutorials run from a server store user work within the storage provided by the end user’s web browser.*).  
The end user does not need to have R installed to run the code chunks. 


### Simple exercise

Here's a simple exercise with an empty code chunk provided for entering the answer.

*Write the R code required to multiply two with three:*

```{r multiplication, exercise=TRUE}

```  
  
  
### Add a solution  
  
We can add a hidden solution button to the exercise chunk.  
  
*Now multiply five with three:*  

```{r multiplication2, exercise=TRUE}

```  
  
```{r multiplication2-solution}
5*3
```
  
  
### Add hints  
  
We can give hints for code execution:  
  
*Now calculate three two the power of five:*  
  
```{r power, exercise=TRUE}

```  

<div id="power-hint">
**Hint:** The operator for an exponent in R is "^". You can find all R operators [here](https://www.datamentor.io/r-programming/operator/)
</div>  
  
  
### Mulitple hints   
  
Note that we can add multiple hints and open them successively:  
  
*Add two vectors, the first one range from 1:2, and the second one from 3:6.* 
  
```{r addvector, exercise = TRUE}

```  
  

```{r addvector-hint-1}
# we can build vectors in R with calling c()
# e.g. this is the vector ranging from 1:2: 
c(1:2) 

# or
c(1,2)

```

```{r addvector-hint-2}
# to add vectors, we can use the plus operator. Note that R functions 
# are vectorized and the length of the shorter vector is therefor recycled: 
c(1:2) + c(3:6)

```



### Preset code 
  

We can already set up some code in the chunks as well, and the user just needs to finnish or modify it.  

*Modify the following code to limit the number of rows printed to 5 of the 'mtcars' data set:*  

```{r print-limit, exercise=TRUE, exercise.eval=TRUE}
mtcars
```

```{r print-limit-hint-1}
# the head function will do the work for you
# note that head generally prints 6 rows
head(mtcars)
```  
  
```{r print-limit-hint-2}
# set the n argument to 5 to print only 5 rows
head(mtcars, n = 5)
``` 
  
  
### Evaluating exercise  
  
Using the `gradethis` package, we can directly evaluate code entered in the code chunks.  
  

* Submit `1+1` to receive a correct grade.

```{r addsimple, exercise = TRUE}

```
 
```{r addsimple-check}
grade_result(
  pass_if(~identical(.result, 2))
)
```



## Topic 2
  
### Quiz

*You can include any number of single or multiple choice questions as a quiz. Use the `question` function to define a question and the `quiz` function for grouping multiple questions together.*

Some questions to verify that you understand the purposes of various base and recommended R packages:

```{r quiz}
quiz(
  question("Which package contains functions for installing other R packages?",
    answer("base"),
    answer("tools"),
    answer("utils", correct = TRUE),
    answer("codetools")
  ),
  question("Which of the R packages listed below are used to create plots?",
    answer("lattice", correct = TRUE),
    answer("tools"),
    answer("stats"),
    answer("grid", correct = TRUE)
  )
)
```

